# Intro to NBGrader
**Jamison McCorrison - 2016**

#### Resources

1. [How it works... Examples of testing](http://nbgrader.readthedocs.org/en/stable/user_guide/04_autograding.html)
    * For autograding assignments, you should have the submitted versions of students’ assignments organized as follows:
        * > `submitted/{student_id}/{assignment_id}/{notebook_id}.ipynb`
    * After running nbgrader autograde, the autograded version of the notebooks will be:
        * > `autograded/{student_id}/{assignment_id}/{notebook_id}.ipynb`
    * Additional auto-generated output files described below.
    * Example student-completed output files ([1](http://nbgrader.readthedocs.org/en/stable/user_guide/submitted/Bitdiddle/Problem%20Set%201/Problem%201.html)) ([2](http://nbgrader.readthedocs.org/en/stable/user_guide/submitted/Hacker/Problem%20Set%201/Problem%201.html))
1. [JupyterHub Config](http://nbgrader.readthedocs.org/en/stable/user_guide/11_jupyterhub_config.html)
2. [nbGrader Manual](https://nbgrader.readthedocs.org/en/stable/)

* * *

# Step 1 : NBGrader (and Jupyter Toolbar) Install

**Install methods:** (local)

```python
%%bash
export PATH="~/anaconda/bin/:$PATH"
cd ~/anaconda
~/anaconda/bin/pip install nbgrader --upgrade
~/anaconda/bin/nbgrader extension install --user
~/anaconda/bin/nbgrader extension activate

# then restart your browser, and/or the jupyter server
```

Output:
```
Requirement already up-to-date: nbgrader in ./lib/python2.7/site-packages
Requirement already up-to-date: notebook in ./lib/python2.7/site-packages (from nbgrader)
Requirement already up-to-date: Flask in ./lib/python2.7/site-packages (from nbgrader)
Requirement already up-to-date: requests in ./lib/python2.7/site-packages (from nbgrader)
Requirement already up-to-date: nbformat in ./lib/python2.7/site-packages (from nbgrader)
Requirement already up-to-date: six in ./lib/python2.7/site-packages (from nbgrader)
Requirement already up-to-date: jupyter-core in ./lib/python2.7/site-packages (from nbgrader)
Requirement already up-to-date: traitlets in ./lib/python2.7/site-packages (from nbgrader)
Requirement already up-to-date: nbconvert in ./lib/python2.7/site-packages (from nbgrader)
Collecting sqlalchemy (from nbgrader)
  Using cached SQLAlchemy-1.0.11.tar.gz
Requirement already up-to-date: python-dateutil in ./lib/python2.7/site-packages (from nbgrader)
Requirement already up-to-date: t
```

# Step 2 : Create Test Cases using All Entry Types
* [Guide](http://nbgrader.readthedocs.org/en/stable/user_guide/02_developing_assignments.html)

### Markdown cells (Text feedback) can be handled as 'manually-graded' or 'read-only':

* This block is read only
* The block below is an example of a manually-graded text blob:

**Who is your favorite TA?** (Olga or Jamison)

<!-- grade_id: 1, points: 0, grade: true, solution: true -->
#### *REPLACE THIS TEXT!*

### Code cells can be set up as...

#### Example input prep and print

```python
%%bash
export variable="ACTGCGTACAGTCAGAT"
echo $variable
```

Output:
```
STRING
```

#### 1. Manually-graded

Replace all instances of "T" to "X" and "a" to "q" and print to screen.

<!-- grade_id: 2, points: 0, grade: true, solution: true -->
```python
%%bash
export variable="ACTGCGTACAGTCAGAT"
echo $variable | tr 'T' 'X' | tr 'a' 'q'
```

Output:
```
ACXGCGXACAGXCAGAX
```

#### 2. Simple example: utograded answer

Count the number of 'C' characters in string $variable.  

<!-- grade_id: 3, points: 0, grade: false, solution: true -->
```python
%%bash
export variable="ACTGCGTACAGTCAGAT"
echo $variable | grep -o 'C'| wc -l
```

Output:
```
4
```

#### 3. Realistic example: Autograder answer + tests

* [Writing good autograder tests](http://nbgrader.readthedocs.org/en/stable/user_guide/02_developing_assignments.html#writing-tests)

**Example from guide above:**  
Write a function that returns a list of numbers, such that xi=i2, for 1≤i≤n. Make sure it handles the case where n<1 by raising a ValueError.

**First set up an 'autograded answer' box for user entry:

<!-- grade_id: 4, points: 0, grade: false, solution: true -->
```python
def squares(n):
    """Compute the squares of numbers from 1 to n, such that the 
    ith element of the returned list equals i^2.
    """
    if n < 1:
        raise ValueError
    s = []
    for i in range(1,n+1):
        s.append(i**2)
    return s
```

**And test:**

```python
squares(1)
```

Output:
```
[1]
```

<!-- grade_id: 4_testsuccess, points: 0, grade: true, solution: false -->
```python
from nose.tools import assert_equal
assert_equal(squares(1), [1])
assert_equal(squares(2), [1, 4])
assert_equal(squares(10), [1, 4, 9, 16, 25, 36, 49, 64, 81, 100])
assert_equal(squares(11), [1, 4, 9, 16, 25, 36, 49, 64, 81, 100, 121])
```

<!-- grade_id: 4_invalidinputcheck, points: 0, grade: true, solution: false -->
```python
from nose.tools import assert_raises
assert_raises(ValueError, squares, 0)
assert_raises(ValueError, squares, -4)
```

#### 4. Read-only

See the code example below:

```python
%%bash
echo "hello world"
```

Output:
```
hello world
```

# Step 3: Autograding (Examples)
* [Guide](http://nbgrader.readthedocs.org/en/stable/user_guide/04_autograding.html)
* Once the autograding is complete, there will be new directories for the autograded versions of the submissions:
    * > `autograded/username/Problem Set 1/Problem 1.ipynb`
* And "Feedback" will be given based on scores from each assert statement
    * > `feedback/username/Problem Set 1/Problem 1.html`
    * [Example](http://nbgrader.readthedocs.org/en/stable/user_guide/feedback/Bitdiddle/Problem%20Set%201/Problem%201.html)
    
**Example** (*Not working because our test case does not follow directory guidelines.*)

```python
%%bash
nbgrader autograde "MyFirstNBGrader"
```

Output:
```
[AutogradeApp | ERROR] No notebooks were matched by 'submitted/*/MyFirstNBGrader'
```

* * *

# The End.